Chemotherapy activates inflammasomes to cause inflammation-associated bone loss

Author Response

The following is the authors’ response to the previous reviews.

Reviewing Editor

We thank you for clarifying several of the questions raised by the reviewers. Since the study has otherwise largely stayed unchanged, we will leave the eLife assessment as “before”:

We respectfully disagree because we addressed all concerns raised by the two reviewers except one (below), which was not satisfactorily answered according to reviewer 1; it has now been addressed (new S3 Fig).

Reviewer #1 (Recommendations For The Authors):

The authors addressed most of my previous comments. However, there is one important point that was not satisfactorily addressed "The band intensities on Western blots in Fig. 4 and Fig. 5 are not quantified, and the numbers of repeats are also not provided" The response that "It is not straightforward to quantify and describe the intensity of the bands of these numerous with different fate outcomes." In the revision, they mentioned at least three repeats were performed. If so, it's not entirely clear why they couldn't quantify the western blots results. Including quantitative data will strengthen the rigor of the findings.

Quantitative data from Fig. 4 and Fig. 5 are now provided as S3 Fig and described in the manuscript (lines 170-175; 184-188).

Changes from the earlier version to the revised version

Abstract

Abstract

Chemotherapy is a widely used treatment for a variety of solid and hematological malignancies. Despite its success in improving the survival rate of cancer patients, chemotherapy causes significant toxicity to multiple organs, including the skeleton, but the underlying mechanisms have yet to be elucidated. Using tumor-free mouse models, which are commonly used to assess direct off-target effects of anti-neoplastic therapies, we found that doxorubicin caused massive bone loss in wild-type mice, a phenotype associated with increased number of osteoclasts, leukopenia, elevated serum levels of danger-associated molecular patterns (DAMPs; e.g., cell-free DNA and ATP) and cytokines (e.g., IL-1β and IL-18). Accordingly, doxorubicin activated the absent in melanoma (AIM2) and NLR family pyrin domain containing 3 (NLRP3) inflammasomes in macrophages and neutrophils, causing inflammatory cell death pyroptosis and NETosis, which correlated with its leukopenic effects. Moreover, the effects of this chemotherapeutic agent on cytokine secretion, cell demise, and bone loss were attenuated to various extent in conditions of AIM2 and/or NLRP3 insufficiency. Thus, we found that inflammasomes are key players in bone loss caused by doxorubicin, a finding that may inspire the development of a tailored adjuvant therapy that preserves the quality of this tissue in patients treated with this class of drugs.

Introduction

The chemotherapeutic drug, doxorubicin, is widely used for the treatment of breast cancer, bladder cancer, lymphomas, and acute lymphocytic leukemia (Smith et al., 2010, Jang et al., 2023, Geffen and Man, 2002). Despite its success in improving the survival rate of cancer patients, doxorubicin causes serious adverse effects, including cardiomyopathy, bone marrow suppression, hair loss, and skeletal manifestations (Cardinale et al., 2015, Tacar et al., 2013, Coleman et al., 2013). Bone complications include osteoporosis, a metabolic disease that is characterized by decreased bone mass and deteriorated microarchitecture, and associated with increased risks for the development of late fractures, and morbidities in the elderly populations (Cauley et al., 2009, Cawthon et al., 2009, Cawthon et al., 2012). In fact, it was reported that 20-50% of geriatric patients (≥65 years) with a hip fracture die within one year of fracture (Coleman et al., 2013). Consistent with the dogma that bone resorption by osteoclasts (OCs) and bone formation by osteoblasts is uncoupled in osteoporotic patients, doxorubicin causes bone loss by promoting osteoclastogenesis while suppressing osteoblastogenesis (Yao et al., 2020, Chai et al., 2014, Rana et al., 2013, Zhou et al., 2020). Increased production of senescence-associated secretory phenotype (SASP), enhanced generation of reactive oxygen species (ROS), and dysregulated autophagy and mitochondrial metabolism are proposed mechanisms of doxorubicin-induced bone pathology (Yao et al., 2020, Park et al., 2022).

Doxorubicin intercalates into DNA thereby impeding the activity of DNA repair enzymes such as topoisomerase II and impairing DNA replication (Robson et al., 1987, Lawrence, 1988, Bonner and Lawrence, 1990, Abe et al., 2022, Tewey et al., 1984). Defective DNA repair ultimately culminates in genomic instability and cell demise, events that can provoke uncontrollable release of intracellular contents such as DNA and various danger-associated molecular patterns (DAMPs). DNA-enriched entities include neutrophil extracellular traps (NETs), web-like structures in which DNA is decorated with peptides, some of which have anti-microbial and inflammatory properties (Komada et al., 2018, Apel et al., 2021). NETs can also propagate inflammation following their engulfment by phagocytes (Boccia et al., 2022, Blayney and Schwartzberg, 2022, Nakazawa et al., 2016, Jeong et al., 2021). Since DNA normally resides in the nucleus and mitochondria, its presence in the cytoplasm is detected by DNA sensors, including absent in melanoma 2 (AIM2), and cyclic guanosine monophosphate-adenosine monophosphate synthase (cGAS), which can trigger immune responses aimed at eliminating the mislocated DNA (Sun et al., 2013, Zhang et al., 2014, Rathinam et al., 2010). Oxidized DNA and various DAMPs can also be sensed by NLRP3 (Shimada et al., 2012, Lu et al., 2014, Xian et al., 2022). Upon recognition of DAMPs or pathogen-associated molecular patterns (PAMPs), AIM2 and NLRP3 assemble protein platforms comprising the adaptor protein apoptosis-associated speck-like protein containing a CARD (ASC) and caspase-1. These protein complexes known as inflammasomes are responsible for the maturation of pro-interleukin-1β (pro-IL-1β) and pro-IL-18 to IL-1β and IL-18, respectively (Guo et al., 2015, Sharma and Kanneganti, 2021, Wang et al., 2021). Inflammasome-comprising caspase-1 also cleaves gasdermin D (GSDMD), generating N-terminal fragments, which form IL-1β- and IL-18-secreting conduits, and cause the inflammatory cell death, pyroptosis (Guo et al., 2015, Sharma and Kanneganti, 2021, Wang et al., 2021). While acute activation of inflammasomes is important for the clearance of the perceived danger and restoration of homeostasis, chronic or excessive stimulation of these safeguard mechanisms can cause disease.

Inflammasomes are involved in skeletal pathophysiology. Gain-of-function mutations of NLRP3 cause skeletal abnormalities in humans (Aksentijevich et al., 2002, Feldmann et al., 2002), which are reproduced to a great extent in knockin mice expressing NLRP3 harboring mutations found in these patients (Bonar et al., 2012, Qu et al., 2015, Snouwaert et al., 2016, Wang et al., 2017). In normal mice, degraded bone matrix components, which are released during bone resorption promote infammasome activation and OC differentiation (Alippe et al., 2017). Age-associated bone loss has also been linked to chronic low-grade inflammation mediated by the NLRP3 inflammasome (Youm et al., 2013). More relevant to this study, radiation, which is also used as an anti-neoplastic therapy, causes bone loss through GSDMD downstream of AIM2 and NLRP3 inflammasomes, but not NLR family caspase recruitment domain containing protein 4 (NLRC4) inflammasome (Xiao et al., 2020). Collectively, the bone phenotypes of genetically or pharmacologically activated inflammasome sensors suggest that the fate of bone cells can be influenced by inflammation-driven by inflammasomes, which are mainly activated in myeloid cells (Brydges et al., 2009, Meng et al., 2009). This view provides a strong rationale for exploring the role of inflammasome pathways in bone loss induced by off-target actions of doxorubicin as this agent causes the death of cancer and bystander normal cells, thereby releasing DAMPs such as ATP and DNA, which activate these pathways.

We used non-tumor-bearing mouse models, which are commonly used to assess off-target outcomes of anti-neoplastic therapies (Park et al., 2022, Borniger et al., 2015, Harrison et al., 1980, Yao et al., 2020) to study bone adverse effects of doxorubicin. We found that doxorubicin caused massive bone loss in wild-type (WT) mice, a phenotype associated with increased number of OCs, leukopenia, and cytokinemia. These outcomes implicated the AIM2 and NLRP3 inflammasomes as they were attenuated upon genetic inactivation of these sensors. Thus, our results show that inflammasomes are key players in bone loss caused by doxorubicin, a finding that may enable the implementation of novel strategies for chemotherapy-related bone complications.

Results

Doxorubicin causes bone loss

To determine the effects of doxorubicin on bone mass, femurs of 10-week-old WT female mice were analyzed by VivaCT before (baseline) and 4 weeks after a single intraperitoneal injection of 5 mg/kg doxorubicin or vehicle. Doxorubicin, but not the vehicle, caused bone loss (S1A-D Fig), findings that were consistent with the recently reported stimulatory and suppressive effects of this drug on bone resorption and formation in WT female mice, respectively (Yao et al., 2020). Doxorubicin also caused a low bone mass phenotype attended by an increased OC number and surface in 10-week-old WT male mice (Fig 1A-E and S1E Fig). Since doxorubicin inflicted bone damage independently of the sex, afterward mechanistic studies focused mainly on male mice and revolved around innate immune responses, which regulate OC-mediated bone resorption in pathological conditions.

Doxorubicin causes bone loss in male mice.

(A, B) Femurs from WT male mice were analyzed by VivaCT before (baseline) and 4 weeks after a single intraperitoneal injection of 5 mg/kg doxorubicin. (A) Cross sections of 3D reconstructions. Scale bar: 500 μm. (B) Bone parameters. (C-E) Femurs from WT male mice, analyzed 4 weeks after a single intraperitoneal injection of vehicle or doxorubicin. Specimens were stained for TRAP activity. (C) Representative images. Scale bar: 500 μm. (D) N.Oc/BS. (E) Oc.S/BS. N=5 mice/group. Data are mean ± SEM. Student t test was used. *P < 0.05; **P < 0.01; ***P < 0.001. BMD, bone mineral density; BV/TV, bone volume/total volume; Dox, doxorubicin; N.Oc/BS, OC number/bone surface; Oc.S/BS, OC surface/bone surface; OC, osteoclast; ns, not significant; Tb.N, trabecular number; Tb.Th, trabecular thickness; Tb.Sp, trabecular separation; WT, wild-type.

Doxorubicin causes cytokinemia, leukopenia, release of DAMPs, and NETosis in vivo

Doxorubicin induces inflammatory responses in patients and experimental models (Wittenburg et al., 2019). Accordingly, WT mice exposed to doxorubicin for 3 days exhibited higher serum levels of IL-1β, IL-18, IL-6, and TNF-α compared to vehicle-injected counterparts (Fig 2A-D). Levels of these inflammatory cytokines inversely correlated with the abundance of white blood cells (WBC; Fig 2E). While doxorubicin lowered the number of circulating lymphocytes and monocytes, the number of neutrophils, the most abundant immune cells in the blood, increased 2 hours post-drug exposure before progressively returning to baseline levels. Consistent with the leukopenic outcome, levels of ATP, which is released by dead cells, were higher in doxorubicin-treated mice compared to vehicle-treated cohorts (Fig 2F). To gain further insight into the mechanisms of leukopenia, we focused on neutrophils, the most abundant leukocytes in blood. These cells exhibit morphological changes such as NET extrusion upon exposure to PAMPs or sterile DAMPs, and eventually undergo NETosis (Komada et al., 2018, Apel et al., 2021). To determine the effects of doxorubicin on NET formation, we measured serum levels of NET components in mice treated with vehicle or this drug for 48 hours. We found that doxorubicin induced NET formation as evidenced by increased levels of citrullinated histone 3 (cit-H3; Fig 2G), myeloperoxidase (MPO; Fig 2H), and cell-free DNA (cfDNA) (Fig 2I, J). Thus, doxorubicin causes cytokinemia, a response that is associated with increased cell death and decreased number of WBCs.

Doxorubicin causes cytokinemia, leukopenia, release of DAMPs, and NETosis in vivo.

Twelve-week-old WT mice were exposed to a single dose of vehicle or 5 mg/kg doxorubicin. Serum samples were harvested 3 days (A-D) or 2 days (F-J) later and analyzed by MSD (IL-1β, IL-6, and TNF-α) or ELISA (IL-18). (E) Blood was collected for cell counts at the indicated time-points after a single dose of 5 mg/kg doxorubicin injection. cfDNA was measured using Qubit (I) or SYBR green (J). Data are mean ± SEM. N=5-12 mice/group. *P < 0.05; **P < 0.01; ***P < 0.001 vs. 0 hour; ##P < 0.01 vs. 2 or 24 hours. Student t test (A-D, F-J) and One-way ANOVA (E) were used. cfDNA, cell free DNA; Cit-H3, citrullinated histone H3; Dox, doxorubicin; IL, interleukin; MPO, myeloperoxidase; WBC, white blood cells.

Doxorubicin activates inflammasome –dependent and – independent pathways, and caused macrophage pyroptosis

The effects of doxorubicin on macrophages have been reported (Chen et al., 2023, Saleh et al., 2021). To test the hypothesis that these cells were implicated in the inflammatory phenotype of mice treated with doxorubicin, we treated bone marrow-derived macrophages (BMMs) with lipopolysaccharide (LPS) for 3 hours to induce the expression of inflammasome components (Wang et al., 2021), then with various concentrations of this drug for 16 hours. Within its reported potent concentrations (1.5 µM – 12 µM) (Saleh et al., 2021, Eljack et al., 2022), doxorubicin did not induce IL-1β secretion, but it significantly caused the release of lactate dehydrogenase (LDH; Fig 3A, B), a marker of cell death as it is released only upon plasma membrane rupture (Wang et al., 2021). Since LDH release was not induced by doxorubicin in a dose-dependent manner, this response may be the result of nonselective cytotoxic actions of this drug. By contrast, doxorubicin promoted IL-1β and LDH release by LPS-primed BMMs in a dose-dependent fashion, with the maximal effect on IL-1β secretion achieved at 6 µM (Fig 3A). Unexpectedly, LPS attenuated LDH release induced by low doxorubicin concentrations (1. 5 and 3 µM) (Fig 3B). To further gain insight into the mechanism of action of doxorubicin, we measured the levels of ATP, which is released by dead cells and activates multiple pathways, including the NLRP3 inflammasome (Karmakar et al., 2016, Carta et al., 2015). ATP levels were higher in the supernatants of doxorubicin-treated BMMs compared to controls, a response that was further enhanced by LPS (Fig 3C). Collectively, these results suggest that BMMs underwent pyroptosis in the presence of LPS and doxorubicin, releasing DAMPs such as ATP.

Doxorubicin activates inflammasomes and causes macrophage pyroptosis.

WT BMMs were left untreated or primed with LPS for 3 hours, then treated with various doxorubicin concentrations for 16 hours. IL-1β (A), LDH (B), and ATP (C) in the conditioned media were measured by ELISA, the cytotoxicity detection Kit or ATP detection Kit, respectively. (D-F) WT BMMs from Asc-citrine mice were primed with 100 ng/ml LPS for 3 hours and treated or not with 15 µM nigericin for 30 minutes or 10 µM doxorubicin for 16 hours. (G) Non-primed cells were also treated with 10 µM doxorubicin for 16 hours. Scale bar: 50 µm. ASC specks were visualized under fluorescence microscopy and quantified using imageJ. (H) Quantitative data. Data are mean ± SEM from experimental triplicates and are representative of at least 2 independent experiments. ***P < 0.001 vs. untreated- or LPS-treated cultures. One-way ANOVA. ASC, apoptosis-associated speck-like protein containing a CARD; ATP, adenosine triphosphate; Dox, doxorubicin; IL-1β, interleukin-1β; LDH, lactate dehydrogenase; LPS, lipopolysaccharide; WT, wild-type.

Earlier version: To reinforce the proposition that doxorubicin activates inflammasomes, we assessed the effects of this drug on the formation of ASC specks, a readout of inflammasome-activated states (Stutz et al., 2013). As anticipated, LPS induced ASC speck formation only in the presence of nigericin, a well-known trigger of NLRP3 inflammasome assembly signals (Fig 3D, E, H). Likewise, doxorubicin induced ASC speck formation only in LPS-primed BMMs (Fig 3F, G, H). Next, we performed immunoblotting to analyze the expression of NLRP3 and AIM2 since these sensors assemble inflammasomes in response to DAMPs such as ATP and DNA, which were released by doxorubicin-damaged cells. We also determined the expression of other key components of these pathways such as caspase-1 and gasdermins. LPS induced the expression of NLRP3, but not AIM2, caspase-1, caspase-3, GSDMD, and GSDME (Fig 4). The active fragments of caspase-1 and GSDMD, p10 and p30, respectively, were mainly detected in cells treated with LPS and doxorubicin (Fig 4 and S2A Fig). GSDMD p10 and GSDME p35, which are proteolytically generated by caspase-3 (S2A, B Fig), were also detected, but their abundance was not affected by LPS (Fig 4). Together, these results suggest that doxorubicin activates both caspase-1 and caspase-3, which cleave GSDMD and GSDME, ultimately, causing pyroptosis and IL-1β release.

Revised version: To reinforce the proposition that doxorubicin activates inflammasomes, we assessed the effects of this drug on the formation of ASC specks, a readout of inflammasome-activated states (Stutz et al., 2013). As anticipated, LPS induced ASC speck formation only in the presence of nigericin, a well-known trigger of NLRP3 inflammasome assembly signals (Fig 3D, E, H). Likewise, doxorubicin induced ASC speck formation only in LPS-primed BMMs (Fig 3F, G, H). Next, we performed immunoblotting to analyze the expression of NLRP3 and AIM2 since these sensors assemble inflammasomes in response to DAMPs such as ATP and DNA, which were released by doxorubicin-damaged cells. We also determined the expression of other key components of these pathways such as caspase-1 and gasdermins. LPS induced the expression of NLRP3, but not AIM2, caspase-1, caspase-3, GSDMD, and GSDME (Fig 4). Levels of caspase-1 (p10) and GSDMD (p30) fragments, which are generated upon inflammasome activation, were higher in cells treated with LPS + doxorubicin compared to doxorubicin alone (Fig 4, S2A Fig and S3A Fig). GSDMD (p10) and GSDME (p35) fragments, which are proteolytically generated by caspase-3 (S2A, B Fig) were also detected, but their abundance was comparable between cells exposed to LPS + doxorubicin and doxorubicin alone (Fig 4 and S3A Fig). Together, these results suggest that doxorubicin activates both caspase-1 and caspase-3, which cleave GSDMD and GSDME, ultimately, causing pyroptosis and IL-1β release.

Doxorubicin activates inflammasome –dependent and –independent pathways in macrophages.

Earlier version: WT BMMs were left untreated or primed with LPS for 3 hours, then treated with various doxorubicin concentrations for 16 hours. Whole cell lysates were analyzed by immunoblotting. Blots are representative of at least 3 independent experiments. AIM2, absent in melanoma 2; cCasp, cleaved caspase; cGSDM, cleaved gasdermin; LPS, lipopolysaccharide; Dox, doxorubicin; WT, wild-type.

Revised version: WT BMMs were left untreated or primed with LPS for 3 hours, then treated with various doxorubicin concentrations for 16 hours. Whole cell lysates were analyzed by immunoblotting. Data are representative of at least 3 independent experiments. AIM2, absent in melanoma 2; cCasp, cleaved caspase; cGSDM, cleaved gasdermin; LPS, lipopolysaccharide; Dox, doxorubicin; WT, wild type.

Doxorubicin activates inflammasome –dependent and – independent pathways, and causes NETosis in vitro

Earlier version: To further support the conclusion that doxorubicin induced the formation of NETs, first, we analyzed the expression of some key players directly or indirectly involved in this process. Incubation of mouse bone marrow neutrophils with LPS for 19 hours resulted in increased NLRP3 expression (Fig 5A). GSDMD, GSDME, caspase-1, and caspase-3 were readily detected in neutrophil lysates (Fig 5A). The abundance of these proteins and their cleaved fragments were unaltered upon cell exposure to LPS, but markedly reduced in doxorubicin-treated cells, likely because of cell death as the levels of β-actin used as loading control were markedly reduced in samples from these cells (Fig 5A). GSDMD p20 and p10 fragments were detected only in neutrophils treated with LPS and doxorubicin. LPS also induced IL-1β secretion, a response that was enhanced by doxorubicin in a dose-dependent manner (Fig 5B). LPS reduced baseline as well as doxorubicin-induced LDH release (Fig 5C). In sum, unchallenged neutrophils released LDH, a response that aligned with the presence of functional caspase-3, GSDMD, and GSDME fragments in these cells. However, LPS was required for optimal NLRP3 expression, caspase-1 activation, and IL-1β production by neutrophils. Although neutrophils secreted higher levels of IL-1β in response to doxorubicin, they were highly sensitive to the cytotoxic effects of this drug.

Revised version: To further support the conclusion that doxorubicin induced the formation of NETs, first, we analyzed the expression of some key players directly or indirectly involved in this process. Incubation of mouse bone marrow neutrophils with LPS for 19 hours resulted in increased NLRP3 expression (Fig 5A). The abundance of caspase-1 (p10) and GSDMD (p30) was indistinguishable between cells exposed to LPS + doxorubicin and doxorubicin alone, likely as the result of cell death as the levels of β-actin used as loading control were markedly reduced in samples from these cells (Fig 5A and S3B Fig). Levels of GSDMD (p10) and GSDME (p35) fragments were also similar between LPS + doxorubicin compared to doxorubicin alone (Fig 5A and S3B Fig). LPS also induced IL-1β secretion, a response that was enhanced by doxorubicin in a dose-dependent manner (Fig 5B). LPS reduced baseline as well as doxorubicin-induced LDH release (Fig 5C). In sum, unchallenged neutrophils released LDH, a response that aligned with the presence of functional caspase-3, GSDMD, and GSDME fragments in these cells. However, LPS was required for optimal NLRP3 expression, caspase-1 activation, and IL-1β production by neutrophils. Although neutrophils secreted higher levels of IL-1β in response to doxorubicin, they were highly sensitive to the cytotoxic effects of this drug.

Doxorubicin activates inflammasome –dependent and –independent pathways in neutrophils.

WT mouse bone marrow neutrophils were left untreated or primed with LPS for 3 hours, then treated with various doxorubicin concentrations for 16 hours. (A) Whole cell lysates were analyzed by immunoblotting. Blots are representative of at least 3 independent experiments. IL-1β (B) and LDH (C) in the conditioned media were measured by ELISA and by the cytotoxicity detection Kit, respectively. Data are mean ± SEM from experimental triplicates and are representative of at least 2 independent experiments. ***P < 0.001 vs. LPS; ##P < 0.01, ###P < 0.001 vs. untreated cultures. One-way ANOVA was used. cCasp, cleaved caspase; cGSDM, cleaved gasdermin; IL-1β, interleukin-1β LDH, lactate dehydrogenase; LPS, lipopolysaccharide; Dox, doxorubicin.

Next, we performed immunofluorescence to visualize NET components in neutrophils treated with vehicle or this drug for 16 hours. Consistent with in vivo results, Cit-H3 and MPO were detected in cells treated with doxorubicin but not with vehicle (Fig 6A). Accordingly, levels of cfDNA were higher in doxorubicin-exposed cultures compared to untreated or cultures treated with LPS (Fig 6B). To determine the biological relevance of cfDNA while modeling the in vivo bone microenvironment, we assessed the impact of degrading cfDNA with DNAse I on IL-1β release by cultured whole bone marrow cells, which comprised various cell types, including neutrophils and macrophages. IL-1β levels were higher in cells treated with LPS and doxorubicin compared to LPS, an outcome that was inhibited by DNAse I (Fig 6C). Thus, doxorubicin promoted IL-1β release, a response that correlated with its effects on NET formation and abundance of cfDNA.

Doxorubicin causes NETosis in vitro.

(A) WT mouse bone marrow neutrophils were left untreated or treated with 10 µM doxorubicin for 16 hours. Cit-H3 and MPO were analyzed by immunofluorescence. Scale bar: 50 μm. Images are representative of at least 3 independent experiments. (B) Neutrophils were left untreated or primed with LPS for 3 hours, then treated with 10 µM doxorubicin for 16 hours. cfDNA in the conditioned medium was extracted and quantified. (C) Neutrophils were left untreated or primed with LPS for 3 hours, then treated with 10 µM doxorubicin and/or DNase I for 16 hours. IL-1β in the conditioned media were measured by ELISA. Data are mean ± SEM from experimental triplicates and are representative of at least 2 independent experiments. *P < 0.05; **P < 0.01; ***P < 0.001 vs. LPS; #P < 0.05, ###P < 0.001 vs. LPS + Dox. One-way ANOVA was used. Dox, doxorubicin; cfDNA, cell free DNA; Cit-H3, citrullinated histone H3; MPO, myeloperoxidase.

AIM2 and NLRP3 inflammasomes are involved in bone-damaging effects of doxorubicin

Earlier version: We hypothesized that doxorubicin induced IL-1β and IL-18 release by activating the AIM2 and NLRP3 inflammasomes as blood levels of their activators (DNA, ATP) were increased n response doxorubicin administration (Fig 2F, I, J). To test this idea, we measured the effects of this drug on IL-1β and LDH release by WT, Aim2-/-, Nlrp3-/-, or Aim2-/-;Nlrp3-/- BMMs and neutrophils. LPS induction of IL-1β and LDH release by BMMs required doxorubicin, a response that was significantly reduced in Aim2 or Nlrp3 deficient cells (Fig 7A, B). LPS induced IL-1β secretion by neutrophils, an outcome that was enhanced by doxorubicin, and comparably attenuated in all mutants (Fig 7C, D). Differences in LDH release were marginal, perhaps because baseline levels of this readout were high, consistent with the short lifespan of these cells in vitro. Because inflammation leads to bone loss, we assessed bone outcomes of 10 weeks old WT or mice globally lacking AIM2 and/or NLRP3, 4 weeks after receiving a single dose of 5 mg/kg doxorubicin or vehicle. Administration of doxorubicin to WT mice caused bone loss associated with increased OC number and surface (Fig 7E-I and S3A-C Fig), consistent with the results shown above (Fig 1). These responses were reduced slightly in Nlrp3 deficient mice, but significantly in Aim2 null mice. Aim2-/-male mice lost bone comparably to Aim2-/-;Nlrp3-/-and caspase-1-/- counterparts (Fig. 7E-G). Similar trends in bone changes were observed in female mice, though Aim2-/- mice were more osteopenic than caspase-1-/- mice (S4A-D Fig). Collectively, our results suggest that the AIM2 and NLRP3 inflammasomes participate to various extent in doxorubicin-bone-damaging effects. They also suggest that inflammasome-independent actions of this drug on bone are not negligible.

Revised version: We hypothesized that doxorubicin induced IL-1β and IL-18 release by activating the AIM2 and NLRP3 inflammasomes as blood levels of their activators (DNA, ATP) were increased n response doxorubicin administration (Fig 2F, I, J). To test this idea, we measured the effects of this drug on IL-1β and LDH release by WT, Aim2-/-, Nlrp3-/-, or Aim2-/-;Nlrp3-/-BMMs and neutrophils. LPS induction of IL-1β and LDH release by BMMs required doxorubicin, a response that was significantly reduced in Aim2 or Nlrp3 deficient cells (Fig 7A, B). LPS induced IL-1β secretion by neutrophils, an outcome that was enhanced by doxorubicin, and comparably attenuated in all mutants (Fig 7C, D). Differences in LDH release were marginal, perhaps because baseline levels of this readout were high, consistent with the short lifespan of these cells in vitro. Because inflammation leads to bone loss, we assessed bone outcomes of 10 weeks old WT or mice globally lacking AIM2 and/or NLRP3, 4 weeks after receiving a single dose of 5 mg/kg doxorubicin or vehicle. Administration of doxorubicin to WT mice caused bone loss associated with increased OC number and surface (Fig 7E-I and S4A-C Fig), consistent with the results shown above (Fig 1). These responses were reduced slightly in Nlrp3 deficient mice, but significantly in Aim2 null mice. Aim2-/- male mice lost bone comparably to Aim2-/-;Nlrp3-/- and caspase-1-/-counterparts (Fig. 7E-G). Similar trends in bone changes were observed in female mice, though Aim2-/- mice were more osteopenic than caspase-1-/- mice (S5A-D Fig). Collectively, our results suggest that the AIM2 and NLRP3 inflammasomes participate to various extent in doxorubicin-bone-damaging effects. They also suggest that inflammasome-independent actions of this drug on bone are not negligible.

AIM2 and NLRP3 inflammasomes are involved in bone-damaging effects of doxorubicin.

WT, Aim2-/-, Nlrp3-/-, Aim2-/-;Nlrp3-/- or Casp1-/- BMMs (A, B) and neutrophils (C, D) were left untreated or primed with LPS for 3 hours, then exposed or not to 10 µM doxorubicin for 16 hours. IL-1β (A, C) and LDH (B, D) in the conditioned media were measured by ELISA and by the cytotoxicity detection Kit, respectively. (E-G) Femurs from male mice were analyzed by VivaCT before (baseline) and 4 weeks after a single intraperitoneal injection of 5 mg/kg doxorubicin. (E) Cross sections of 3D reconstructions. Scale bar: 500 μm. (F) BV/TV changes. (G) BMD changes. (H, I) Femurs harvested from different genotypes male mice were analyzed 4 weeks after a single intraperitoneal injection of vehicle or doxorubicin. Specimens were stained for TRAP activity. (H) N.Oc/BS. (I) Oc.S/BS. Data are mean ± SEM from experimental triplicates and are representative of at least 2 independent experiments (A-D); n=5-8 mice/group (E-I). Data are mean ± SEM. *P < 0.05; **P < 0.01; ***P < 0.001 vs. control, LPS or vehicle; ##P < 0.01, ###P < 0.001 vs. other genotypes or WT treated with Dox. Two-way ANOVA (A-D, H-I) and One-way ANOVA (F-G) were used. AIM2, absent in melanoma 2; BMD bone mineral density; BV/TV, bone volume/total volume; casp1, caspase-1; Cont, control; Dox, doxorubicin; IL-1β, interleukin-1β; LDH, lactate dehydrogenase; LPS, lipopolysaccharide; N.Oc/BS, OC number/bone surface; Oc.S/BS; OC, osteoclast; WT, wild-type.

Discussion

We found that the AIM2 inflammasome and the NLRP3 inflammasome to a lesser extent, played an important role in bone-damaging effects of doxorubicin. The comparable bone phenotype of Aim2-/-;Nlrp3-/- and caspase-1-/-mice suggested that the AIM2 and NLRP3 inflammasomes were the primary mediators of doxorubicin actions. Because doxorubicin activates several pathways, some of which interact or overlap with inflammasome functions (e.g., senescence factors), the remaining bone loss in compound mutant mice was expected. The interplay among these pathways may have accounted for the sex differences in the bone outcomes of inflammasome insufficiency as residual doxorubicin-induced bone loss was higher in Aim2-/- and Aim2-/-;Nlrp3-/- female mice than in male counterparts. Sexual dimorphic actions of inflammasomes were not unprecedented as uneven severity of atherosclerosis was found in male and female Nlrp3-deficient mice with gonadal insufficiency (Chen et al., 2020) and sex-dependent differential activation of AIM2 and NLRP3 inflammasomes in macrophages from systemic lupus erythematous had been reported (Yang et al., 2015). Although the impacts of doxorubicin on bone pathology are complex, including its direct actions on bone cells (Yao et al., 2020, Chai et al., 2014, Rana et al., 2013, Zhou et al., 2020), our investigation focused on immune cells, and found that the inactivation of inflammasomes is sufficient to attenuate the drug’s bone-damaging effects.

Doxorubicin causes neutropenia, lymphopenia, and anemia in patients (Gatti et al., 2018, Boccia et al., 2022). Consistent with the clinical situation, administration of doxorubicin to mice caused leukopenia, which correlated with lymphocytopenia and monocytopenia, and was associated with fluctuations in neutrophil counts. Since neutrophils were highly sensitive to the cytotoxic effects of doxorubicin, the transient neutrophilic effects of this drug in mice may be the result of emergency granulopoiesis, a physiological response that is rapidly triggered to restore adequate neutrophil number. This leukopenic outcome was consistent with increased serum levels of cell death-associated DAMPS (ATP and cfDNA) and our results showing that doxorubicin activated the effectors of apoptosis, pyroptosis, and NETosis, including caspase-1, caspase-3, GSDMD, and GSDME. Since pyroptosis and NETosis promote inflammation and immune responses, we argued that they accounted for the cytokinemic effects of doxorubicin. Other studies, however, found that doxorubicin reduced NET formation in cancer models and by human neutrophils in vitro (Lu et al., 2021, Khan et al., 2019). This discrepancy may be due to differences in the experimental models and cell context-dependent actions of doxorubicin. Other limitations of our study include its focus on: i) macrophages and neutrophils while oversighting lymphocytes or even other cells bone microenvironment such as mesenchymal and adipocytes whose fate is affected by this drug (Fan et al., 2018, Fan et al., 2017, Wang et al., 2012, Buttiglieri et al., 2011), which were also targeted by doxorubicin; ii) immune cells without assessing the direct effects of doxorubicin on bone cells, as noted above, and iii) the use of the tumor-free model as immune responses can differ significantly in the absence or presence of cancer cells. Despite these limitations, our findings point to a novel mechanism of action for doxorubicin in bone.

DNA accumulates in the cytoplasm as the result of genomic instability, damaged mitochondria, or lysed intracellular pathogens. Extracellular DNA from pathogens, NETotic, or pyroptotic cells can be internalized and culminate in the cytoplasm. In either case, sensors such as AIM2 detect mislocated DNA in the cytoplasm and triggers inflammatory responses (Boccia et al., 2022, Blayney and Schwartzberg, 2022, Nakazawa et al., 2016, Jeong et al., 2021). This view was consistent with our data showing a correlation between the levels of extracellular DNA and IL-1β as well as by the inhibition of IL-1β secretion by DNAse I. We also found that doxorubicin activated the NLRP3 inflammasome and induced the release of ATP, a well-known activator of the NLRP3 inflammasome. Whether doxorubicin activated the NLRP3 inflammasome directly by perturbing the plasma membrane or indirectly via generation of secondary signals such as ATP remained unclear.

Earlier version: By showing that inflammasomes are key players in bone loss caused by doxorubicin, this work advances our knowledge on potential mechanisms of action of this drug on this tissue (S5 Fig). This insight is difficult to get in clinical situations because this chemotherapeutic is employed not alone but in conjunction with other medications (Zimny, 1988, Svendsen et al., 2017, Müller et al., 2020).

Revised version: By showing that inflammasomes are key players in bone loss caused by doxorubicin, this work advances our knowledge on potential mechanisms of action of this drug on this tissue (S6 Fig). This insight is difficult to get in clinical situations because this chemotherapeutic is employed not alone but in conjunction with other medications (Zimny, 1988, Svendsen et al., 2017, Müller et al., 2020).

Materials and Methods

Animals

Earlier version: Casp1-/- were kindly provided by Dr. Thirumala-Devi Kanneganti (St. Jude Children’s Research Hospital). WT, Asc-citrine, Aim2-/-, and Nlrp3-/- mice were purchased from The Jackson Laboratory (Sacramento, CA). Aim2-/- mice and Nlrp3-/- mice were intercrossed to generate Aim2-/-;Nlrp3-/- mice. All mice were on the C57BL6J background and mouse genotyping was performed by PCR. All procedures were approved by the Institutional Animal Care and Use Committee (IACUC) of Washington University School of Medicine in St. Louis. All experiments were performed in accordance with the relevant guidelines and regulations described in the IACUC-approved protocol 22-0335.

Revised version: Casp1-/- were kindly provided by Dr. Thirumala-Devi Kanneganti (St. Jude Children’s Research Hospital). WT, Asc-citrine, Aim2-/-, and Nlrp3-/- mice were purchased from The Jackson Laboratory (Sacramento, CA). Aim2-/- mice and Nlrp3-/- mice were intercrossed to generate Aim2-

Added in the revised version: /-;Nlrp3-/- mice. All mice were on the C57BL6J background and mouse genotyping was performed by PCR. All procedures were approved by the Institutional Animal Care and Use Committee (IACUC) of Washington University School of Medicine in St. Louis. All experiments were performed in accordance with the relevant guidelines and regulations described in the IACUC-approved protocol 22-0335.

Doxorubicin administration and VivaCT analysis

The femurs of 10-week-old mice were analyzed by VivaCT 2 weeks before (baseline) and 4 weeks after a single intraperitoneal injection (i.p.) of 5 mg/kg doxorubicin (Sigma-Aldrich, MO) formulated in H2O at 1 mg/ml or vehicle. For bone analysis, mice were anesthetized with isofluorane and trabecular volume in the distal femoral metaphysis of the right leg was measured using VivaCT 40 (Scanco Medical AG, Zurich, Switzerland) set at 70 kVp, 114 μA, and 20 μm resolution as previously described (Yao et al., 2020). For the trabecular bone compartment, contours were traced on the inside of the cortical shell using 2D images of the femoral metaphysis. The end of the growth plate region was used as a landmark to establish a consistent location for starting analysis, and the next 50 slices were analyzed. The following trabecular parameters are reported for all Viva CT experiments: bone volume over total volume (BV/TV), trabecular number (Tb.N), trabecular thickness (Tb.Th), trabecular separation (Tb.Sp), and volumetric bone mineral density (vBMD).

Histomorphometry

For static histomorphometry, the femurs were fixed in 10% neutral buffered formalin overnight and decalcified in 14% (wt/vol) EDTA, pH 7.2, for 10–14 days at room temperature. Fixed femurs were embedded in paraffin, sectioned at 5-μm thicknesses, and mounted on glass slides. The sections were stained with tartrate-resistant acidic phosphatase (TRAP) as described previously (Wang et al., 2018). For dynamic histomorphometry, mice were i.p. injected with 10 mg/kg calcein green (Sigma-Aldrich, MO) and 6 days later with 50 mg/kg alizarin red (Sigma-Aldrich, MO). Mice were euthanized 2 days after the second injection. The tibias were collected and fixed in 10% neutral buffered formalin overnight, embedded in methyl methacrylate, and sectioned at 7–10 μm. Images were taken using a Nanozoomer 2.0 HT whole slide scanner (Hamamatsu Photonics, Shizuoka, Japan) at 20X magnification. Bioquant Osteo software (v18.2.6; Bioquant Image Analysis Corp, TN) was used for image analysis. Measurements of dynamic bone histomorphometry were calculated from fluorochrome double labels at the endocortical surfaces as previously described (Xiao et al., 2020).

Serum assays

Blood was collected by cardiac puncture and was allowed to clot at room temperature. Serum obtained after centrifugation at 2,000 x g for 10 minutes was used for various assays. Cytokine and chemokine levels were measured by V-PLEX Plus Proinflammatory Panel 1 Mouse Kit (Meso Scale Diagnostics, MD), except IL-18, which was analyzed by enzyme linked immunosorbent assay (ELISA) kit (Sigma-Aldrich, MO). The levels of Citrine-H3 (Cit-H3) and myeloperoxidase (MPO) were determined by ELISA kits (Abcam, MA & Cayman, MI).

Peripheral blood analysis

Mouse blood was collected by cardiac puncture in the EDTA-containing tubes. Complete blood counts were performed by the Washington University School of Medicine as previously described (Wang et al., 2017).

Cell cultures

Murine primary bone marrow-derived macrophages (BMMs) were obtained by culturing mouse bone marrow cells in culture media containing a 1:10 dilution of supernatant from the fibroblastic cell line CMG 14-12 as a source of macrophage colony-stimulating factor, a mitogenic factor for BMMs, for 4-5 days in a 15-cm dish as previously described (Takeshita et al., 2000, Wang et al., 2020). Briefly, nonadherent cells were removed by vigorous washes with PBS, and adherent BMMs were detached with trypsin-EDTA and cultured in culture media containing a 1:10 dilution of CMG for various experiments. Murine primary neutrophils were isolated by collecting bone marrow cells and subsequently over a discontinuous Percoll (Sigma Aldrich, MO) gradient as described previously (Sun et al., 2022). Briefly, all bone marrow cells from femurs and tibias were washed by PBS and then resuspended in 2 ml PBS. Cell suspension was gently layered on top of gradient (72% Percoll, 64% Percoll, 52% Percoll) and centrifuged at 1545× g for 30 minutes at room temperature. After carefully discarding the top two cell layers, the third layer containing neutrophils was transferred to a clean 15 ml tube. Cells were washed and counted, then plated at a density of 1 × 105 cells/well in 96-well plate or 5 × 106 cells/well in 6-well plate for 1 hour followed by various experiments. For all in vitro experiments except otherwise specified, BMMs were plated at 2×104 cells per well on a 96-well plate or 106 cells per well on a 6-well plate overnight. Neutrophils were plated at 105 cells per well on a 96-well plate or 5 × 106 cells per well on a 6-well plate for one hour prior to treatment. BMMs and neutrophils were primed with 100 ng/ml LPS (Sigma-Aldrich, MO) for 3 hours, then with different concentrations of doxorubicin (Sigma-Aldrich, MO) as indicated for 16 hours. Conditioned media were collected for the analysis of IL-1β and lactate dehydrogenase (LDH). Cell lysates were collected for protein expression analysis by Western blot as described below.

Western blot analysis

Cell extracts were prepared by lysing cells with RIPA buffer (50 mM Tris, 150 mM NaCl, 1 mM EDTA, 0.5% NaDOAc, 0.1% SDS, and 1.0% NP-40) plus phosphatase and protease inhibitors (GenDEPOT, TX). Protein concentrations were determined by the Bio-Rad Laboratories method (Bio-Rad, CA), and equal amounts of proteins were subjected to SDS-PAGE gels (12% or 15%) as previously described (Wang et al., 2021). Proteins were transferred onto nitrocellulose membranes and incubated with antibodies against GSDMD (1;1,000; Abcam, MA), GSDME (1;1,000; Abcam, MA), caspase-1 (1;1,000; Abcam, MA), caspase-3 (1:1,000; Cell Signaling Technologies, MA), NLRP3 (1:1,000; AdipoGen,CA), AIM2 (1:1000; Cell Signaling Technologies, MA) or β-actin (1:2,000; Santa Cruz Biotechnology, TX) overnight at 4°C followed by incubation for 1 hour with secondary goat anti-mouse IRDye 800 (Thermo Fisher Scientific, MA) or goat anti–rabbit Alexa Fluor 680 (Thermo Fisher Scientific, MA), respectively. The results were visualized using the Odyssey infrared imaging system (LI-COR Biosciences, NE).

LDH assay and IL-1β ELISA

Cell death was assessed by the release of LDH in conditioned medium using LDH Cytotoxicity Detection Kit (TaKaRa, CA). IL-1β levels in conditioned media were measured by an ELISA kit (eBiosciences, NY).

ASC specks assay

Asc-citrine-WT BMMs were plated at 104 cells per well on a 16-well glass plate overnight. Cells were primed with LPS for 3 hours followed by 15 μM nigericin (AdipoGen, CA) for 30 minutes or the indicated doxorubicin concentrations for 16 hours. Cells were washed with PBS, fixed with 4% paraformaldehyde buffer for 10 minutes at room temperature, then counterstained with Fluoro-gel II containing DAPI (Fluoro-Gel, Fisher Scientific Intl INC, PA). Asc-citrine photographs were taken using ZEISS microscopy (Carl Zeiss Industrial Metrology, MN). Quantification of ASC specks was carried out using ImageJ.

Immunofluorescence

Isolated neutrophils were plated at 105 cells per well on a 16-well glass plate for 1 hour. Cells were primed with LPS for 3 hours, treated with doxorubicin for 16 hours, washed with PBS, and fixed with 4% paraformaldehyde buffer for 10 minutes at room temperature. Cells were permeabilized with 0.2% Triton in PBS for 20 minutes, blocked with 0.2% Triton, 1% BSA, and CD61 antibody (1:1000; Alexa Fluor 647, BD, NJ) in PBS for 30 minutes, and were incubated with Cit-H3 antibody (1:1,000; Cayman, MI) and MPO (1;1,000; Abcam, MA) overnight at 4°C in blocking buffer, followed by incubation with secondary antibody (Alexa Fluor 594, 1:2,000; Life Technologies, CA) for 30 minutes. Cells were counterstained with Fluoro-gel II containing DAPI (Fluoro-Gel, Fisher Scientific Intl INC, PA). Immunostaining images were taken using a Leica inverted microscope with a TCS SPEII confocal module and processed using LAS X software (Leica Microsystems Inc, IL).

cfDNA assay

cfDNA in the conditioned cell culture medium was extracted using NucleoSpin Gel and PCR Clean-up kit (Takara, Duren, Germany), and quantified with Nanodrop (Thermo Fisher Scientific, MA). cfDNA in the serum was purified and measured using Qubit by Washington University School of Medicine Hope Center DNA/RNA purification Core or using SYBR Green (Thermo Fisher Scientific, MA) as previously described (Goldshtein et al., 2009, Villalba-Campos et al., 2016).

ATP assay

ATP levels in conditioned media and serum were measured by RealTime-Glo Extracellular ATP Assay kit (Promega, Madison, WI).

Statistical analysis

Statistical analysis was performed using the Student’s t test, one-way ANOVA with Tukey’s multiple comparisons test, or two-way ANOVA with Tukey’s multiple comparisons test, Dunnett’s multiple comparisons test, or Sidak’s multiple comparisons test using the GraphPad Prism 9.0 Software. Values are expressed as mean ± SEM. *p < 0.05 was considered statistically significant.

Supporting information